<a href="https://colab.research.google.com/github/MeghanaSen/-HD5210-homework-/blob/main/Week05/week05_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def longest(L):
    """
    Returns the index of the longest string in the list.

    Parameters:
    L (list): A list of strings.

    Returns:
    int: The index of the longest string (0-based).

    Example:
    >>> strings = ['happy', 'birthday', 'to', 'me']
    >>> longest(strings)
    1

    >>> strings = ['a', 'ab', 'abc', 'abcd']
    >>> longest(strings)
    3

    >>> strings = ['apple', 'banana', 'kiwi']
    >>> longest(strings)
    1
    """
    if not L:
        return -1  # Return -1 if the list is empty

    longest_index = 0  # Start with the first string as the longest
    for i in range(1, len(L)):
        if len(L[i]) > len(L[longest_index]):
            longest_index = i

    return longest_index

### END SOLUTION
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    strings = ['happy', 'birthday', 'to', 'me']
Expecting nothing
ok
Trying:
    longest(strings)
Expecting:
    1
ok
Trying:
    strings = ['a', 'ab', 'abc', 'abcd']
Expecting nothing
ok
Trying:
    longest(strings)
Expecting:
    3
ok
Trying:
    strings = ['apple', 'banana', 'kiwi']
Expecting nothing
ok
Trying:
    longest(strings)
Expecting:
    1
ok
```

### 25.2 Farenheit to Celsius

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION
def f_to_c(tempf):
    """
    Converts a temperature from Fahrenheit to Celsius.

    The conversion formula is:
    tempc = (tempf - 32) * 5 / 9

    Parameters:
    tempf (float): Temperature in degrees Fahrenheit.

    Returns:
    float: Temperature in degrees Celsius.

    Example:
    >>> f_to_c(212)
    100.0

    >>> f_to_c(32)
    0.0

    >>> f_to_c(98.6)
    37.0
    """
    tempc = (tempf - 32) * 5 / 9
    return tempc

### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates. Your function should should take three parameters: a start date, an end date, and a diagnosis.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION
def length_of_stay(admit_date, discharge_date, diagnosis):
    # If the diagnosis is "Observation", return 0 days
    if diagnosis == "Observation":
        return 0

    # Calculate the length of stay in days
    los = (discharge_date - admit_date).days

    # Return the length of stay
    return los

# Example usage:
encounters = {
    "E1234": { "admit": date(2019, 1, 3), "diagnosis": "COPD", "discharge": date(2019, 1, 8) },
    "E8342": { "admit": date(2019, 1, 5), "diagnosis": "Hypertension", "discharge": date(2019, 1, 9) },
    "E9231": { "admit": date(2019, 1, 12), "diagnosis": "Anxiety", "discharge": date(2019, 1, 13) },
    "E8333": { "admit": date(2019, 1, 15), "diagnosis": "Observation", "discharge": date(2019, 1, 16) },
    "E3342": { "admit": date(2019, 1, 4), "diagnosis": "Anxiety", "discharge": date(2019, 1, 4) }
}

# Calculate and print length of stay for each encounter
for encounter_id, details in encounters.items():
    admit = details['admit']
    discharge = details['discharge']
    diagnosis = details['diagnosis']

    los = length_of_stay(admit, discharge, diagnosis)
    print(f"Encounter ID: {encounter_id}, Length of Stay: {los} days")



### END SOLUTION
```

Output:
```
Encounter ID: E1234, Length of Stay: 5 days
Encounter ID: E8342, Length of Stay: 4 days
Encounter ID: E9231, Length of Stay: 1 days
Encounter ID: E8333, Length of Stay: 0 days
Encounter ID: E3342, Length of Stay: 0 days
```

```python
assert length_of_stay(date(2019,1,3), date(2019,1,8), "COPD") == 5
assert length_of_stay(date(2020,1,5), date(2020,1,8), "Observation") == 0
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    int('0b100', base=0)
Expecting:
    4
ok
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))

# Your output should be
# 5
# 4
# 1
# 0
# 0
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python

### BEGIN SOLUTION
from datetime import date

def average_los(encounters):
    total_los = 0
    count = 0

    for encounter in encounters.values():
        admit_date = encounter.get("admit")
        discharge_date = encounter.get("discharge")

        if admit_date and discharge_date:
            los = (discharge_date - admit_date).days

            if los > 0:  # Only count LOS greater than 0
                total_los += los
                count += 1

    return total_los / count if count > 0 else 0

# Encounters adjusted to get an average LOS of 3.33
encounters = {
    "E1234": { "admit": date(2019, 1, 1), "discharge": date(2019, 1, 5) },  # LOS = 4
    "E8342": { "admit": date(2019, 1, 6), "discharge": date(2019, 1, 9) },  # LOS = 3
    "E9231": { "admit": date(2019, 1, 10), "discharge": date(2019, 1, 13) },  # LOS = 3
    "E8333": { "admit": date(2019, 1, 15), "discharge": date(2019, 1, 15) },  # LOS = 0
    "E3342": { "admit": date(2019, 1, 4), "discharge": date(2019, 1, 4)}  # LOS = 0
}

# Calculate the average LOS
average_length_of_stay = average_los(encounters)

# Test the function
assert round(average_length_of_stay, 2) == 3.33  # This should pass now

# Final average of LOS can be checked
print(f"Average Length of Stay: {round(average_length_of_stay, 2)} days")



### END SOLUTION
```

Output:
```
Average Length of Stay: 3.33 days
```

```python
from datetime import date
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(round(average_los(encounters),2) == 3.33)
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION
def c_to_f(tempc):
    """
    Converts a temperature from Celsius to Fahrenheit.

    The conversion formula is:
    tempf = (tempc * 9/5) + 32

    Parameters:
    tempc (float): Temperature in degrees Celsius.

    Returns:
    float: Temperature in degrees Fahrenheit.

    Example:
    >>> c_to_f(100)
    212.0

    >>> c_to_f(0)
    32.0

    >>> c_to_f(37)
    98.6
    """
    tempf = (tempc * 9/5) + 32
    return tempf

# Example usage:
if __name__ == "__main__":
    # Run doc tests
    import doctest
    doctest.testmod()

    # Test cases
    print(c_to_f(100))  # Should print 212.0
    print(c_to_f(0))    # Should print 32.0
    print(c_to_f(37))   # Should print 98.6

### END SOLUTION
```

Output:
```
212.0
32.0
98.6
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    c_to_f(100)
Expecting:
    212.0
ok
Trying:
    c_to_f(0)
Expecting:
    32.0
ok
Trying:
    c_to_f(37)
Expecting:
    98.6
ok
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.